```python
# Assignment
```

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"

    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    # raise NotImplementedError()
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re


def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids.""" 

    return re.findall("[A-Z]", simple_string)
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re


def grades():
    with open("assets/grades.txt", "r") as file:
        grades = file.read()
    return re.findall(r":\sB", grades)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re

def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    pattern = r"""
    (?P<host>.*)         # Host IP
    (\ -\ )              # Separator 1
    (?P<user_name>\w*|-)   # The User name doesn't have a space
    (\s\[)               # space and brackets for the time delimitation
    (?P<time>.*)         # time of requisition
    (\]\ ")              # separatior of Requisition
    (?P<request>\w*\ .*)  # Request Specs only grab POST
    ("\w*\n)                   # end of line
    """
    result = []
    for item in re.finditer(pattern, logdata, re.VERBOSE):
        result.append(item.groupdict())
    return result
```

```python
assert len(logs()) == 979
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[39], line 1
----> 1 assert len(logs()) == 979

AssertionError:
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item = {
    "host": "146.204.224.152",
    "user_name": "feest6811",
    "time": "21/Jun/2019:15:45:24 -0700",
    "request": "POST /incentivize HTTP/1.1",
}
assert (
    one_item in logs()
), "Sorry, this item should be in the log results, check your formating"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[38], line 1
----> 1 assert len(logs()) == 979
      3 one_item = {
      4     "host": "146.204.224.152",
      5     "user_name": "feest6811",
      6     "time": "21/Jun/2019:15:45:24 -0700",
      7     "request": "POST /incentivize HTTP/1.1",
      8 }
      9 assert (
     10     one_item in logs()
     11 ), "Sorry, this item should be in the log results, check your formating"

AssertionError:
```

```python

```

```python

```

```python

```

```python

```